---
title: "Introduction to data"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE,message=FALSE}
library(BIOL710AdvBiometry)
library(learnr)
library(learnrhash)
library(tidyverse)
library(gradethis)
tutorial_options(
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )
knitr::opts_chunk$set(echo = FALSE)
# stuff to load in secret
data(nycbirds)
la_bird_obs <- nycbirds %>%
  filter(dest == "LA")
sf_bird_obs <- nycbirds %>%
  filter(dest == "SF")
sf_feb_bird_obs <- nycbirds %>%
  filter(dest == "SF", month == 2)
```

## Getting started

Some define statistics as the field that focuses on turning information into
knowledge. The first step in that process is to summarize and describe the raw
information -- the data. In this lab we explore migratory birds, specifically an artificial random
sample of migratory birds that were spotted in
New York area locations in 2013, tagged, and whose routes were followed through to their destination. We will generate simple graphical and numerical 
summaries of data on these birds and explore variation in timing. Since this is a large 
data set, along the way you'll also learn the indispensable skills of data 
processing and subsetting. Please note that these data are not real and are simply for demonstrating data exploration in R. 

### Load packages

In this lab, we will explore and visualize the data using the **tidyverse** suite of 
packages. 

Let's load the package.

```{r load-packages, exercise = T}
library(tidyverse)
```


## The data

First, we'll view the `nycbirds` data frame. Type the following in your console
to load the data:

```{r load-data, exercise = T}
data(nycbirds)
```

The data set `nycbirds` that shows up in your workspace is a *data matrix*, 
with each row representing an *observation* and each column representing a 
*variable*. R calls this data format a **data frame**, which is a term that will 
be used throughout the labs. For this data set, each *observation* is a single migratory bird observation.

To view the names of the variables, type the command

```{r names, exercise = T}
names(nycbirds)
```

This returns the names of the variables in this data frame. The **codebook**
(description of the variables) can be accessed by pulling up the help file:

```{r help, exercise = T}
?nycbirds
```

One of the variables refers to the taxonomic order of the bird observed (bird_obs), which 
is coded according to the following system.

- `birds`: Two letter bird abbreviation.
    + `DU`:           Ducks
    + `FL`:           Flamingos 
    + `DO`:           Doves
    + `CU`:           Cuckoos
    + `NI`:           Nightjars
    + `SW`:           Swifts
    + `CR`:           Cranes
    + `SH`:           Shore Birds
    + `TR`:           Tropicbirds
    + `LO`:           Loons
    + `TU`:           Tube-nosed Swimmers
    + `ST`:           Storks
    + `CO`:           Cormorants
    + `PE`:           Pelicans
    + `HA`:           Hawks # 
    + `OW`:           Owls
    + `TG`:           Trogons
    + `MO`:           Motmots
    + `WO`:           Woodpeckers
    + `FA`:           Falcons
    + `PB`:           Perching Birds


We can use `glimpse` to take a quick peek at your data to 
understand its contents better.

```{r glimpse, exercise = T}
glimpse(nycbirds)
```


## Making plots to explore the data

We use the `ggplot()` function to build plots. 
If you run the plotting code in your console, you should see the plot appear under the *Plots* tab of the lower right panel of RStudio. If you run the chunk, the plot will appear beneath the code chunk.
 
With `ggplot()`:

- The first argument is always the dataset. 
- Next, you provide the variables from the dataset to be assigned to `aes`thetic elements of the plot, e.g. the x and the y axes. 
- Finally, you use another layer, separated by a `+` to specify the `geom`etric object for the plot. Since we want to scatterplot, we use `geom_point()`.

For instance, if you wanted to visualize the above plot using a line graph, you would replace `geom_point()` with `geom_line()`.

### Departure delays

Let's start by examing the distribution of migration shifts of all migratory bird observations with a 
histogram.

```{r hist-dep-delay, exercise = T}
ggplot(data = nycbirds, aes(x = mig_shift)) +
  geom_histogram()
```

This function says to plot the `mig_shift` variable from the `nycbirds` data 
frame on the x-axis. It also defines a `geom` (short for geometric object), 
which describes the type of plot you will produce. 

Histograms are generally a very good way to see the shape of a single 
distribution of numerical data, but that shape can change depending on how the 
data is split between the different bins. You can easily define the binwidth you 
want to use:

```{r hist-dep-delay-bins, exercise = T}
ggplot(data = nycbirds, aes(x = mig_shift)) +
  geom_histogram(binwidth = 15)
ggplot(data = nycbirds, aes(x = mig_shift)) +
  geom_histogram(binwidth = 150)
```

Look carefully at these three histograms. How do they compare? 

```{r bandwidth}
question("Are features revealed in one that are obscured in another?",
    answer("No, they look the same."),
    answer("Yes, more detail is shown with a binwidth of 15.", correct = T),
    answer("Yes, more detail is shown with a binwidth of 150."),
    allow_retry = TRUE,
    random_answer_order = F
  )
```

## Piping

The `%>%` operator is called the **piping** operator. It takes the output of the previous expression and pipes it into the first argument of the function in the following one. 
To continue our analogy with mathematical functions, `x %>% f(y)` is equivalent to `f(x, y)`.

If you want to visualize only on delays of migratory bird observations headed to Los Angeles, you need 
to first `filter` the data for migratory bird observations with that destination (`dest == "LA"`)
and then make a histogram of the departure delays of only those migratory bird observations. Note that a common mistake is not using the double equal signs. Remember one equal sign assigns "la" to the variable `dest` while two equal signs asks the question "does each value in `dest` match 'la'?".

```{r la-bird_obs-hist, exercise = T}
la_bird_obs <- nycbirds %>%
  filter(dest == "LA")

ggplot(data = la_bird_obs, aes(x = mig_shift)) +
  geom_histogram()
```

**A note on piping:** Note that we can read these two lines of code as the following: *"Take the `nycbirds` dataset and **pipe** it into the `filter` function. 
Filter the `nycbirds` data set so that we only have rows where the destination is la.

With all of that in mind, let's decipher these two commands (OK, so it might look like four lines, but
the first two physical lines of code are actually part of the same command. It's 
common to add a break to a new line after `%>%` to help readability).

- Command 1: Take the `nycbirds` data frame, `filter` for migratory bird observations headed to LA, and 
save the result as a new data frame called `la_bird_obs`.
    + `==` means "if it's equal to".
    + `la` is in quotation marks since it is a character string.
- Command 2: Basically the same `ggplot` call from earlier for making a histogram, 
except that it uses the smaller data frame for migratory bird observations headed to LA instead of all
migratory bird observations.

<div id="boxedtext">
**Logical operators: ** Filtering for certain observations (e.g. migratory bird observations from a 
particular location) is often of interest in data frames where we might want to 
examine observations with certain characteristics separately from the rest of 
the data. To do so, you can use the `filter` function and a series of 
**logical operators**. The most commonly used logical operators for data 
analysis are as follows:

- `==` means "equal to"
- `!=` means "not equal to"
- `>` or `<` means "greater than" or "less than"
- `>=` or `<=` means "greater than or equal to" or "less than or equal to"
</div>

You can also obtain numerical summaries for these migratory bird observations:

```{r la-bird_obs-summ, exercise = T}
la_bird_obs %>%
  summarise(mean_dd   = mean(mig_shift), 
            median_dd = median(mig_shift), 
            n         = n())
```

Note that in the `summarise` function you created a list of three different
numerical summaries that you were interested in. The names of these elements are 
user defined, like `mean_dd`, `median_dd`, `n`, and you can customize these names 
as you like (just don't use spaces in your names). Calculating these summary 
statistics also requires that you know the function calls. Note that `n()` reports
the sample size. Note you can also use `summarize()` instead of `summarise()` if you prefer that spelling.

<div id="boxedtext">
**Summary statistics: ** Some useful function calls for summary statistics for a 
single numerical variable are as follows:

- `mean`
- `median`
- `sd`
- `var`
- `IQR`
- `min`
- `max`

Note that each of these functions takes a single vector as an argument and 
returns a single value. 
</div>

## More filtering

You can also filter based on multiple criteria. Suppose you are interested in
migratory bird observations headed to San Francisco (SF) in February:

```{r sf-feb-bird_obs, exercise = T}
sf_feb_bird_obs <- nycbirds %>%
  filter(dest == "SF", month == 2)
```

Note that you can separate the conditions using commas (or "&") if you want migratory bird observations that
are both headed to sf **and** in February. If you are interested in either
bird_obs headed to sf **or** in February, you can use the `|` instead of the comma. (This is located above the backslash.)

Create a new data frame that includes migratory bird observations headed to SF in June, 
    and save this data frame as `sf_june_bird_obs`. 
    
```{r filterfeb, exercise = TRUE}
___ <- ___ %>% ___(dest == ___, month == ___)
```

```{r filterfeb-solution}
sf_june_bird_obs <- nycbirds %>% filter(dest == "SF", month == 6)
```

```{r filterfeb-check}
# check code
gradethis::grade_code()
```
    
How many migratory bird observations meet these criteria? **Hint:** How many rows are in `sf_june_bird_obs`? To learn what a function does and its arguments that are available to you, use the "Help" tab in the bottom right. Try this for `dim` and `nrow`. 

```{r dimbird_obs-setup}
sf_june_bird_obs <- nycbirds %>% filter(dest == "SF", month == 6)

```

```{r dimbird_obs, exercise = TRUE}
dim(___)

nrow(___)
```


```{r dimbird_obs-solution}
dim(sf_june_bird_obs)

nrow(sf_june_bird_obs)
```

```{r dimbird_obs-check}
# check code
gradethis::grade_code()
```


## Summary statistics

Another useful technique is quickly calculating summary statistics for various 
groups in your data frame. For example, we can modify the above command using 
the `group_by` function to get the same summary stats for each origin location:

```{r summary-custom-list-origin, exercise = T}
sf_feb_bird_obs %>%
  group_by(origin) %>%
  summarise(median_dd = median(mig_shift), iqr_dd = IQR(mig_shift), n_bird_obs = n())
```

Here, we first grouped the data by `origin` and then calculated the summary 
statistics.

Calculate the median and interquartile range for `arr_shift`s of migratory bird observations in the `sf_feb_bird_obs` data frame, grouped by bird type.  


```{r summarybird, exercise = TRUE}
___ %>%
  ___(___) %>%
  ___(median_dd = ___(___), iqr_dd = ___(___))
```

```{r summarybird-hint-1}
sf_feb_bird_obs %>%
  ___(___) %>%
  summarise(median_dd = ___(arr_shift), iqr_dd = ___(arr_shift))
```

```{r summarybird-solution}
sf_feb_bird_obs %>%
  group_by(carrier) %>%
  summarise(median_dd = median(arr_shift), iqr_dd = IQR(arr_shift))

```

```{r summarybird-check}
# check code
gradethis::grade_code()
```



```{r whichbird}
question("Which carrier has the most variable arrival delays?",
    answer("FL"),
    answer("CU"),
    answer("NI and ST", correct = T),
    answer("PE"),
    allow_retry = TRUE,
    random_answer_order = F
  )
```



### Migration shifts for New York area locations

Suppose you want to test whether different New York area locations are experiencing different migration shifts. For this, any observed migration more than 5 days earlier than average will be considered
"early".

In order to rank the migration start date shifts by area, 
you can:

- first classify each observed migration as "average" or "early",
- then group observed migration by origin location,
- then calculate on time migration start rates for each origin location,
- and finally arrange the locations in descending order for migration shift.

Let's start with classifying each observed migration as "average" or "early" by
creating a new variable with the `mutate` function.

```{r dep-class, exercise = T}
nycbirds <- nycbirds %>%
  mutate(class = ifelse(mig_shift > -5, "average", "early"))
```

The first argument in the `mutate` function is the name of the new variable
we want to create, in this case `class`. Then if `mig_shift > -5`, we classify 
the migratory bird observations as `"average"` and `"early"` if not, i.e. if the migratory bird observations is early 
by 5 or more days

Note that we are also overwriting the `nycbirds` data frame with the new 
version of this data frame that includes the new `class` variable.

We can handle all of the remaining steps in one code chunk:

```{r locrate-setup}
nycbirds <- nycbirds %>%
  mutate(class = ifelse(mig_shift > -5, "average", "early))
```

```{r locrate, exercise = TRUE}
nycbirds %>%
  group_by(origin) %>%
  summarise(loc_rate = sum(class == "average") / n()) %>%
  arrange(desc(loc_rate))
```


***

## Submit checkpoint

```{r context="server"}
learnrhash::encoder_logic(strip_output = T)
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = div(strong("Submit your hash in the form below."), br(), br()),
  ui_after  = learnrhash::iframe_ui(
    src = "https://docs.google.com/forms/d/e/1FAIpQLSfKAR9MMxfAR8EWsCi-U2Ow4SjsejLlAUba14FBsch-gox0Gg/viewform?usp=sf_link",
    width="900px", height= "1000px"
  )
)
```



* * *

<a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons License" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/88x31.png" /></a><br />This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a> then was adapted and learnr-ified by Sara Stoudt, and further adapted and biologified by Jenna Ekwealor.
