---
title: "Introduction to data"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE,message=FALSE}
library(learnr)
library(learnrhash)
library(tidyverse)
library(gradethis)
library(openintro)
tutorial_options(
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )
knitr::opts_chunk$set(echo = FALSE)

la_bird_obss <- nycbirds %>%
  filter(dest == "la")

sf_feb_bird_obss <- nycbirds %>%
  filter(dest == "sf", month == 2)
```

## Getting started

Some define statistics as the field that focuses on turning information into
knowledge. The first step in that process is to summarize and describe the raw
information -- the data. In this lab we explore migratory birds, specifically a random
sample of migratory birds that were spotted in
New York in 2013. We will generate simple graphical and numerical 
summaries of data on these birds and explore variation in timing. Since this is a large 
data set, along the way you'll also learn the indispensable skills of data 
processing and subsetting.

### Load packages

In this lab, we will explore and visualize the data using the **tidyverse** suite of 
packages. The data can be found in the companion package for OpenIntro labs, **openintro**.

Let's load the packages.

```{r load-packages, exercise = T}
library(tidyverse)
library(openintro)
```


## The data

First, we'll view the `nycbirds` data frame. Type the following in your console
to load the data:

```{r load-data, exercise = T}
data(nycbirds)
```

The data set `nycbirds` that shows up in your workspace is a *data matrix*, 
with each row representing an *observation* and each column representing a 
*variable*. R calls this data format a **data frame**, which is a term that will 
be used throughout the labs. For this data set, each *observation* is a single bird_obs.

To view the names of the variables, type the command

```{r names, exercise = T}
names(nycbirds)
```

This returns the names of the variables in this data frame. The **codebook**
(description of the variables) can be accessed by pulling up the help file:

```{r help, exercise = T}
?nycbirds
```

One of the variables refers to the order of the bird observed (bird_obs), which 
is coded according to the following system.

- `birds`: Two letter bird abbreviation.
    + `DU`:           Ducks
    + `FL`:           Flamingos 
    + `DO`:           Doves
    + `CU`:           Cuckoos
    + `NI`:           Nightjars
    + `SW`:           Swifts
    + `CR`:           Cranes
    + `SH`:           Shore Birds
    + `TR`:           Tropicbirds
    + `LO`:           Loons
    + `TU`:           Tube-nosed Swimmers
    + `ST`:           Storks
    + `CO`:           Cormorants
    + `PE`:           Pelicans
    + `HA`:           Hawks
    + `OW`:           Owls
    + `TG`:           Trogons
    + `MO`:           Motmots
    + `WO`:           Woodpeckers
    + `FA`:           Falcons
    + `PB`:           Perching Birds


We can use `glimpse` to take a quick peek at your data to 
understand its contents better.

```{r glimpse, exercise = T}
glimpse(nycbirds)
```


## Making plots to explore the data

We use the `ggplot()` function to build plots. 
If you run the plotting code in your console, you should see the plot appear under the *Plots* tab of the lower right panel of RStudio. If you run the chunk, the plot will appear beneath the code chunk.
 
With `ggplot()`:

- The first argument is always the dataset. 
- Next, you provide the variables from the dataset to be assigned to `aes`thetic elements of the plot, e.g. the x and the y axes. 
- Finally, you use another layer, separated by a `+` to specify the `geom`etric object for the plot. Since we want to scatterplot, we use `geom_point()`.

For instance, if you wanted to visualize the above plot using a line graph, you would replace `geom_point()` with `geom_line()`.

### Departure delays

Let's start by examing the distribution of departure delays of all bird_obss with a 
histogram.

```{r hist-dep-delay, exercise = T}
ggplot(data = nycbirds, aes(x = dep_delay)) +
  geom_histogram()
```

This function says to plot the `dep_delay` variable from the `nycbirds` data 
frame on the x-axis. It also defines a `geom` (short for geometric object), 
which describes the type of plot you will produce. 

Histograms are generally a very good way to see the shape of a single 
distribution of numerical data, but that shape can change depending on how the 
data is split between the different bins. You can easily define the binwidth you 
want to use:

```{r hist-dep-delay-bins, exercise = T}
ggplot(data = nycbirds, aes(x = dep_delay)) +
  geom_histogram(binwidth = 15)
ggplot(data = nycbirds, aes(x = dep_delay)) +
  geom_histogram(binwidth = 150)
```

Look carefully at these three histograms. How do they compare? 

```{r bandwidth}
question("Are features revealed in one that are obscured in another?",
    answer("No, they look the same."),
    answer("Yes, more detail is shown with a binwidth of 15.", correct = T),
    answer("Yes, more detail is shown with a binwidth of 150."),
    allow_retry = TRUE,
    random_answer_order = F
  )
```

## Piping

The `%>%` operator is called the **piping** operator. It takes the output of the previous expression and pipes it into the first argument of the function in the following one. 
To continue our analogy with mathematical functions, `x %>% f(y)` is equivalent to `f(x, y)`.

If you want to visualize only on delays of bird_obss headed to Los Angeles, you need 
to first `filter` the data for bird_obss with that destination (`dest == "la"`)
and then make a histogram of the departure delays of only those bird_obss. Note that a common mistake is not using the double equal signs. Remember one equal sign assigns "la" to the variable `dest` while two equal signs asks the question "does each value in `dest` match 'la'?".

```{r la-bird_obss-hist, exercise = T}
la_bird_obss <- nycbirds %>%
  filter(dest == "la")

ggplot(data = la_bird_obss, aes(x = dep_delay)) +
  geom_histogram()
```

**A note on piping:** Note that we can read these two lines of code as the following: *"Take the `nycbirds` dataset and **pipe** it into the `filter` function. 
Filter the `nycbirds` data set so that we only have rows where the destination is la.

With all of that in mind, let's decipher these two commands (OK, so it might look like four lines, but
the first two physical lines of code are actually part of the same command. It's 
common to add a break to a new line after `%>%` to help readability).

- Command 1: Take the `nycbirds` data frame, `filter` for bird_obss headed to la, and 
save the result as a new data frame called `la_bird_obss`.
    + `==` means "if it's equal to".
    + `la` is in quotation marks since it is a character string.
- Command 2: Basically the same `ggplot` call from earlier for making a histogram, 
except that it uses the smaller data frame for bird_obss headed to la instead of all
bird_obss.

<div id="boxedtext">
**Logical operators: ** Filtering for certain observations (e.g. bird_obss from a 
particular airport) is often of interest in data frames where we might want to 
examine observations with certain characteristics separately from the rest of 
the data. To do so, you can use the `filter` function and a series of 
**logical operators**. The most commonly used logical operators for data 
analysis are as follows:

- `==` means "equal to"
- `!=` means "not equal to"
- `>` or `<` means "greater than" or "less than"
- `>=` or `<=` means "greater than or equal to" or "less than or equal to"
</div>

You can also obtain numerical summaries for these bird_obss:

```{r la-bird_obss-summ, exercise = T}
la_bird_obss %>%
  summarise(mean_dd   = mean(dep_delay), 
            median_dd = median(dep_delay), 
            n         = n())
```

Note that in the `summarise` function you created a list of three different
numerical summaries that you were interested in. The names of these elements are 
user defined, like `mean_dd`, `median_dd`, `n`, and you can customize these names 
as you like (just don't use spaces in your names). Calculating these summary 
statistics also requires that you know the function calls. Note that `n()` reports
the sample size. Note you can also use `summarize()` instead of `summarise()` if you prefer that spelling.

<div id="boxedtext">
**Summary statistics: ** Some useful function calls for summary statistics for a 
single numerical variable are as follows:

- `mean`
- `median`
- `sd`
- `var`
- `IQR`
- `min`
- `max`

Note that each of these functions takes a single vector as an argument and 
returns a single value. 
</div>

## More filtering

You can also filter based on multiple criteria. Suppose you are interested in
bird_obs headed to San Francisco (sf) in February:

```{r sf-feb-bird_obss, exercise = T}
sf_feb_bird_obss <- nycbirds %>%
  filter(dest == "sf", month == 2)
```

Note that you can separate the conditions using commas (or "&") if you want bird_obss that
are both headed to sf **and** in February. If you are interested in either
bird_obss headed to sf **or** in February, you can use the `|` instead of the comma. (This is located above the backslash.)

Create a new data frame that includes bird_obss headed to sf in June, 
    and save this data frame as `sf_june_bird_obss`. 
    
```{r filterfeb, exercise = TRUE}
___ <- ___ %>% ___(dest == ___, month == ___)
```

```{r filterfeb-solution}
sf_june_bird_obss <- nycbirds %>% filter(dest == "sf", month == 6)
```

```{r filterfeb-check}
# check code
gradethis::grade_code()
```
    
How many bird_obss meet these criteria? **Hint:** How many rows are in `sf_june_bird_obss`? To learn what a function does and its arguments that are available to you, use the "Help" tab in the bottom right. Try this for `dim` and `nrow`. 

```{r dimbird_obss-setup}
sf_june_bird_obss <- nycbirds %>% filter(dest == "sf", month == 6)

```

```{r dimbird_obss, exercise = TRUE}
dim(___)

nrow(___)
```


```{r dimbird_obss-solution}
dim(sf_june_bird_obss)

nrow(sf_june_bird_obss)
```

```{r dimbird_obss-check}
# check code
gradethis::grade_code()
```


## Summary statistics

Another useful technique is quickly calculating summary statistics for various 
groups in your data frame. For example, we can modify the above command using 
the `group_by` function to get the same summary stats for each origin airport:

```{r summary-custom-list-origin, exercise = T}
sf_feb_bird_obss %>%
  group_by(origin) %>%
  summarise(median_dd = median(dep_delay), iqr_dd = IQR(dep_delay), n_bird_obss = n())
```

Here, we first grouped the data by `origin` and then calculated the summary 
statistics.

Calculate the median and interquartile range for `arr_delay`s of bird_obss in the `sf_feb_bird_obss` data frame, grouped by carrier. 


```{r summarycarrier, exercise = TRUE}
___ %>%
  ___(___) %>%
  ___(median_dd = ___(___), iqr_dd = ___(___))
```

```{r summarycarrier-hint-1}
sf_feb_bird_obss %>%
  ___(___) %>%
  summarise(median_dd = ___(arr_delay), iqr_dd = ___(arr_delay))
```

```{r summarycarrier-solution}
sf_feb_bird_obss %>%
  group_by(carrier) %>%
  summarise(median_dd = median(arr_delay), iqr_dd = IQR(arr_delay))

```

```{r summarycarrier-check}
# check code
gradethis::grade_code()
```



```{r whichcarrier}
question("Which carrier has the most variable arrival delays?",
    answer("AA"),
    answer("B6"),
    answer("DL and UA", correct = T),
    answer("VX"),
    allow_retry = TRUE,
    random_answer_order = F
  )
```



### On time departure rate for NYC airports

Suppose you will be flying out of NYC and want to know which of the 
three major NYC airports has the best on time departure rate of departing bird_obss. 
Also supposed that for you, a bird_obs that is delayed for less than 5 minutes is 
basically "on time."" You consider any bird_obs delayed for 5 minutes of more to be 
"delayed".

In order to determine which airport has the best on time departure rate, 
you can 

- first classify each bird_obs as "on time" or "delayed",
- then group bird_obss by origin airport,
- then calculate on time departure rates for each origin airport,
- and finally arrange the airports in descending order for on time departure
percentage.

Let's start with classifying each bird_obs as "on time" or "delayed" by
creating a new variable with the `mutate` function.

```{r dep-type, exercise = T}
nycbirds <- nycbirds %>%
  mutate(dep_type = ifelse(dep_delay < 5, "on time", "delayed"))
```

The first argument in the `mutate` function is the name of the new variable
we want to create, in this case `dep_type`. Then if `dep_delay < 5`, we classify 
the bird_obs as `"on time"` and `"delayed"` if not, i.e. if the bird_obs is delayed 
for 5 or more minutes.

Note that we are also overwriting the `nycbirds` data frame with the new 
version of this data frame that includes the new `dep_type` variable.

We can handle all of the remaining steps in one code chunk:

```{r otdeprate-setup}
nycbirds <- nycbirds %>%
  mutate(dep_type = ifelse(dep_delay < 5, "on time", "delayed"))
```

```{r otdeprate, exercise = TRUE}
nycbirds %>%
  group_by(origin) %>%
  summarise(ot_dep_rate = sum(dep_type == "on time") / n()) %>%
  arrange(desc(ot_dep_rate))
```

    
```{r select-airport}
question("If you were selecting an airport simply based on the on time departure 
    percentage, which NYC airport would you choose to fly out of?",
    answer("LGA", correct =T),
    answer("JFK"),
    answer("EWR"),
    answer("None of the above."),
    allow_retry = TRUE,
    random_answer_order = F
  )
```



***

## Submit checkpoint

```{r context="server"}
learnrhash::encoder_logic(strip_output = T)
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = div(strong("Submit your hash in the form below."), br(), br()),
  ui_after  = learnrhash::iframe_ui(
    src = "https://docs.google.com/forms/d/e/1FAIpQLSfKAR9MMxfAR8EWsCi-U2Ow4SjsejLlAUba14FBsch-gox0Gg/viewform?usp=sf_link",
    width="900px", height= "1000px"
  )
)
```

## Creating a reproducible lab report

For the rest of this lab you will R Markdown to create a reproducible lab report. 
In RStudio, go to New File -> R Markdown... Then, choose From Template and then choose `Lab Report` from the list of templates. Make sure to name the document appropriately and pick a location for the file where you know how to find it on your computer.

See the following video describing how to get started with creating these 
reports for this lab, and all future labs:

[**Basic R Markdown with an OpenIntro Lab**](https://www.youtube.com/watch?v=Pdc368lS2hk)
<iframe width="560" height="315" src="https://www.youtube.com/embed/Pdc368lS2hk" frameborder="0" allowfullscreen></iframe>

**Note: This video was made using `oilabs`. We will NOT be using `oilabs`. Be sure to pick the Lab Report template that goes with `openintro` as shown in screenshot above. Make sure you have `library(openintro)` in your first chunk not `library(oilabs)`.**

## Questions for Lab Report

Answer in an Rmd file based on the lab report template.

1.  Suppose you really dislike departure delays and you want to schedule 
    your travel in a month that minimizes your potential departure delay leaving 
    NYC. One option is to choose the month with the lowest mean departure delay.
    Another option is to choose the month with the lowest median departure delay. 
    What are the pros and cons of these two choices?
    
2. The `nycbirds` data frame is a massive trove of information. Pick one of the following questions and tell me which R functions you would use to investigate the question.

- How delayed were bird_obss that were headed to Los Angeles?
- How do departure delays vary by month?
- Which of the three major NYC airports has the best on time percentage for 
departing bird_obss?

**Hint** Example: Which month would you expect to have the highest average delay departing from an NYC airport?

Let's think about how you could answer this question:

- First, calculate monthly averages for departure delays. With the new language
you are learning, you could
    + `group_by` months, then
    + `summarise` mean departure delays.
- Then, you could to `arrange` these average delays in `desc`ending order

```{r mean-dep-delay-months, eval = F, echo = T}
nycbirds %>%
  group_by(month) %>%
  summarise(mean_dd = mean(dep_delay)) %>%
  arrange(desc(mean_dd))
```

3.  Describe the distribution (center, variation, range, skewness/symmetry) of the **arrival** delays of these bird_obss. Would the mean or median be a more appropriate measure of center? Would the standard deviation or the interquartile range be a more appropriate measure of variability. Why?

```{r}
sf_feb_bird_obss <- nycbirds %>%
  filter(dest == "sf", month == 2)

ggplot(data = sf_feb_bird_obss, aes(x = arr_delay)) +
  geom_histogram(binwidth = 5)
```

4.  Load the `nycbirds` dataset and replicate the following plot. **Hint:** The data frame plotted only
    contains bird_obss from American Airlines, Delta Airlines, and United
    Airlines, and the points are `color`ed by `carrier`. Once you replicate
    the plot, determine (roughly) what the cutoff point is for departure
    delays where you can still expect to get to your destination on time.

```{r plot-to-replicate, echo=FALSE, fig.show="asis", fig.width=7, fig.height=4}
dl_aa_ua <- nycbirds %>%
  filter(carrier == "AA" | carrier == "DL" | carrier == "UA")
ggplot(data = dl_aa_ua, aes(x = dep_delay, y = arr_delay, color = carrier)) +
  geom_point()
```

You may want to refer to the [ggplot cheatsheet](https://rstudio.com/wp-content/uploads/2016/11/ggplot2-cheatsheet-2.1.pdf) and the [dplyr cheatsheet](https://rstudio.com/wp-content/uploads/2015/02/data-wrangling-cheatsheet.pdf).


## Deliverables

Make sure you have submitted your hash to the Google Form for the first part of the lab.

When you are finished editing your Markdown document click the "Knit" button and choose "Knit to HTML" in the top left corner of RStudio. This will run all of your code and create a formatted document of the output. If you get an error, it means something in your Markdown file isn't right, either an error in code or some error in formatting. Call me into your breakout room, and we will troubleshoot.

Submit your Markdown document and knitted file to [Moodle](https://moodle.smith.edu/) as:

LastName-LastName-LastName-L-02.Rmd  (add a third last name if applicable)

LastName-LastName-LastName-L-02.html

*Due*: Monday,  (beginning of class time, Eastern time)


* * *

<a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons License" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/88x31.png" /></a><br />This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a> then was adapted and learnr-ified by Sara Stoudt, and further adapted and biologified by Jenna Ekwealor.
