---
title: "Comparing Two Means: t-test"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE,message=FALSE}
library(BIOL710)
library(learnr)
library(learnrhash)
library(tidyverse)
library(gradethis)
library(infer)
tutorial_options(
  exercise.checker = gradethis::grade_learnr
)
knitr::opts_chunk$set(echo = FALSE)

# Load and filter data
control_cells <- heat_shock_response %>%
  dplyr::filter(condition == "control")
heat_stressed_cells <- heat_shock_response %>%
  dplyr::filter(condition == "heat_stressed")
```

## Getting Started

In this lab, you'll explore the t-test, a statistical method used to compare two independent group means and assess if they are significantly different from each other.

### Load Packages

We will analyze data using the `tidyverse` and `infer` packages. Load the necessary packages below:

```{r load-packages, exercise = T}
library(tidyverse)
library(infer)
```

### The Data

We will analyze experimental data measuring heat shock protein (HSP) expression levels in control versus heat-stressed mammalian cells.

```{r viewdata, exercise = TRUE}
glimpse(heat_shock_response)
head(heat_shock_response)
```

## Summary Statistics

The goal is to compare the mean expression levels of HSPs between control and heat-stressed cells.

### Step 1: Compute Summary Statistics

Calculate the mean and standard deviation of HSP expression for both conditions.

```{r summarystats1, exercise = TRUE}
control_mean <- mean(control_cells$HSP_expression)
control_sd <- sd(control_cells$HSP_expression)
```


```{r summarystats2, exercise = TRUE}
heat_stressed_mean <- ___(heat_stressed_cells$___)
heat_stressed_sd <- ___(heat_stressed_cells$___)
```

```{r summarystats2-hint}
# solution code
heat_stressed_mean <- mean(heat_stressed_cells$HSP_expression)
heat_stressed_sd <- sd(heat_stressed_cells$HSP_expression)
```

```{r summarystats2-solution}
# solution code
heat_stressed_mean <- mean(heat_stressed_cells$HSP_expression)
heat_stressed_sd <- sd(heat_stressed_cells$HSP_expression)
```

```{r summarystats2-check}
# check code
gradethis::grade_code()
```

### Step 2: Visualizing the Distributions

Create histograms to visualize the distributions of HSP expression levels in the control condition.

```{r histogram1, exercise = TRUE}
ggplot(control_cells, aes(x = HSP_expression)) + geom_histogram()
```

Now your turn! Create a histogram to visualize the distributions of HSP expression levels in the heat-stressed cells.

```{r histogram2, exercise = TRUE}
ggplot(___, aes(x = ___)) + geom____()
```

```{r histogram2-hint}
# solution code
ggplot(heat_stressed_cells, aes(x = HSP_expression)) + geom_histogram()
```

```{r histogram2-solution}
# solution code
ggplot(heat_stressed_cells, aes(x = HSP_expression)) + geom_histogram()
```

```{r histogram2-check}
# check code
gradethis::grade_code()
```

## Comparing Means with a t-test

### Step 3: Conducting a t-test

We will now perform an independent t-test to determine if the mean HSP expression differs significantly between the control and heat-stressed cells.

```{r ttest, exercise = TRUE}
t_test_result <- t.test(control_cells$___, heat_stressed_cells$___)
print(t_test_result)
```

```{r ttest-hint}
# solution
t_test_result <- t.test(control_cells$HSP_expression, heat_stressed_cells$HSP_expression)
print(t_test_result)
```

```{r ttest-solution}
# solution
t_test_result <- t.test(control_cells$HSP_expression, heat_stressed_cells$HSP_expression)
print(t_test_result)
```

### Step 4: Interpretation

Review the t-test output and determine if the p-value suggests a significant difference.

What does the p-value indicate about the difference in HSP expression levels between the two conditions?

*Write your interpretation as a comment in the code chunk below. Make sure to "run code" to submit your response!

```{r interpretation, exercise = TRUE}
# insert intreptation here
```

## Visualization

Using the infer package, we can further visualize the hypothesis testing process.

```{r vizualize, exercise = TRUE}
heat_shock_response %>%
  specify(response = HSP_expression, explanatory = condition) %>%
  hypothesize(null = "independence") %>%
  calculate(stat = "t")
```

## Submit checkpoint

```{r context="server"}
learnrhash::encoder_logic(strip_output = T)
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = div(strong("Submit your hash in the form below."), br(), br()),
  ui_after  = learnrhash::iframe_ui(
    src = "https://docs.google.com/forms/d/e/1FAIpQLSf0FKvlOHfchtSxwIpD9Jo65sDlTjeokodsO1vP64rLcVYZ2A/viewform?usp=header",
    width="900px", height= "1000px"
  )
)
```


